<!DOCTYPE html>
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<html lang="en-US">
<head xmlns:og="http://ogp.me/ns#" xmlns:book="https://ogp.me/ns/book#">
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<title>SB Writing your own Exercises</title>
<meta name="Keywords" content="Authored in PreTeXt">
<meta name="viewport" content="width=device-width, initial-scale=1.0">
<meta property="og:type" content="book">
<meta property="book:title" content="Using eBooks with Runestone Interactive">
<meta property="book:author" content="Brad Miller, PhD. Professor Emeritus Luther College">
<script>var runestoneMathReady = new Promise((resolve) => window.rsMathReady = resolve);
window.MathJax = {
  tex: {
    inlineMath: [['\\(','\\)']],
    tags: "none",
    tagSide: "right",
    tagIndent: ".8em",
    packages: {'[+]': ['base', 'extpfeil', 'ams', 'amscd', 'color', 'newcommand', 'knowl']}
  },
  options: {
    ignoreHtmlClass: "tex2jax_ignore|ignore-math",
    processHtmlClass: "process-math",
  },
  chtml: {
    scale: 0.98,
    mtextInheritFont: true
  },
  loader: {
    load: ['input/asciimath', '[tex]/extpfeil', '[tex]/amscd', '[tex]/color', '[tex]/newcommand', '[pretext]/mathjaxknowl3.js'],
    paths: {pretext: "https://pretextbook.org/js/lib"},
  },
  startup: {
    pageReady() {
      return MathJax.startup.defaultPageReady().then(function () {
      console.log("in ready function");
      rsMathReady();
      }
    )}
  },
};
</script><script src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-chtml.js"></script><link href="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/themes/prism.css" rel="stylesheet">
<script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/components/prism-core.min.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/autoloader/prism-autoloader.min.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/lunr.js/2.3.9/lunr.min.js" integrity="sha512-4xUl/d6D6THrAnXAwGajXkoWaeMNwEKK4iNfq5DotEbLPAfk6FSxSP3ydNxqDgCw1c/0Z1Jg6L8h2j+++9BZmg==" crossorigin="anonymous" referrerpolicy="no-referrer"></script><script src="lunr-pretext-search-index.js" async=""></script><script src="https://pretextbook.org/js/0.3/pretext_search.js"></script><link href="https://pretextbook.org/css/0.7/pretext_search.css" rel="stylesheet" type="text/css">
<script>js_version = 0.3</script><script src="https://pretextbook.org/js/lib/jquery.min.js"></script><script src="https://pretextbook.org/js/lib/jquery.sticky.js"></script><script src="https://pretextbook.org/js/lib/jquery.espy.min.js"></script><script src="https://pretextbook.org/js/0.3/pretext.js"></script><script>miniversion=0.1</script><script src="https://pretextbook.org/js/0.3/pretext_add_on.js?x=1"></script><script src="https://pretextbook.org/js/0.3/user_preferences.js"></script><script src="https://pretextbook.org/js/lib/knowl.js"></script><!--knowl.js code controls Sage Cells within knowls--><script>sagecellEvalName='Evaluate (Sage)';
</script><link rel="preconnect" href="https://fonts.googleapis.com">
<link rel="preconnect" href="https://fonts.gstatic.com" crossorigin="">
<link href="https://fonts.googleapis.com/css2?family=Inconsolata:wght@400;700&amp;family=Noto+Serif:ital,wght@0,400;0,700;1,400;1,700&amp;family=Tinos:ital,wght@0,400;0,700;1,400;1,700&amp;display=swap" rel="stylesheet">
<link href="https://fonts.cdnfonts.com/css/dejavu-serif" rel="stylesheet">
<link href="https://fonts.googleapis.com/css2?family=Roboto+Serif:opsz,wdth,wght@8..144,50..150,100..900&amp;display=swap" rel="stylesheet">
<link href="https://fonts.googleapis.com/css2?family=Open+Sans:wdth,wght@75..100,300..800&amp;display=swap" rel="stylesheet">
<link href="https://pretextbook.org/css/0.7/pretext.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/pretext_add_on.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/shell_crc.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/banner_crc.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/navbar_crc.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/toc_crc.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/knowls_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/style_oscarlevin.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/colors_blue_red.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.7/setcolors.css" rel="stylesheet" type="text/css">
<!--** eBookCongig is necessary to configure interactive       **-->
<!--** Runestone components to run locally in reader's browser **-->
<!--** No external communication:                              **-->
<!--**     log level is 0, Runestone Services are disabled     **-->
<script type="text/javascript">
eBookConfig = {};
eBookConfig.useRunestoneServices = false;
eBookConfig.host = 'http://127.0.0.1:8000';
eBookConfig.course = 'PTX_Course_Title_Here';
eBookConfig.basecourse = 'PTX_Base_Course';
eBookConfig.isLoggedIn = false;
eBookConfig.email = '';
eBookConfig.isInstructor = false;
eBookConfig.logLevel = 0;
eBookConfig.username = '';
eBookConfig.readings = null;
eBookConfig.activities = null;
eBookConfig.downloadsEnabled = false;
eBookConfig.allow_pairs = false;
eBookConfig.enableScratchAC = false;
eBookConfig.build_info = "";
eBookConfig.python3 = null;
eBookConfig.runestone_version = '7.0.3';
eBookConfig.jobehost = '';
eBookConfig.proxyuri_runs = '';
eBookConfig.proxyuri_files = '';
eBookConfig.enable_chatcodes =  false;
</script>
<!--*** Runestone Services ***-->
<script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.3/prefix-runtime.743dc745e5a9b334.bundle.js"></script><script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.3/prefix-347.193e36c07abd0197.bundle.js"></script><script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.3/prefix-runestone.a6d252836fb5c953.bundle.js"></script><link rel="stylesheet" type="text/css" href="https://runestone.academy/cdn/runestone/7.0.3/prefix-347.f9add1ca35d5ad93.css">
<link rel="stylesheet" type="text/css" href="https://runestone.academy/cdn/runestone/7.0.3/prefix-runestone.931faef5361cf6e9.css">
</head>
<body id="instructorguide" class="pretext book ignore-math">
<a class="assistive" href="#ptx-content">Skip to main content</a><header id="ptx-masthead" class="ptx-masthead"><div class="ptx-banner">
<a id="logo-link" class="logo-link" href="https://runestone.academy" target="_blank"><img src="external/RAIcon.png" alt="Logo image"></a><div class="title-container">
<h1 class="heading"><a href="InstructorGuide.html"><span class="title">Using eBooks with Runestone Interactive:</span> <span class="subtitle">The PreTeXt Interactive Edition</span></a></h1>
<p class="byline">Brad Miller, PhD. Professor Emeritus Luther College</p>
</div>
<div id="searchresultsplaceholder" class="searchresultsplaceholder" style="display: none">
<button id="closesearchresults" class="closesearchresults" onclick="document.getElementById('searchresultsplaceholder').style.display = 'none'; return false;">x</button><h2>Search Results: <span id="searchterms" class="searchterms"></span>
</h2>
<div id="searchempty" class="searchempty"><span>No results.</span></div>
<ol id="searchresults" class="searchresults"></ol>
</div>
</div></header><nav id="ptx-navbar" class="ptx-navbar navbar"><button class="toc-toggle button" aria-label="Show or hide table of contents"><span class="icon">☰</span><span class="name">Contents</span></button><a class="index-button button" href="index-1.html" title="Index"><span class="name">Index</span></a><button id="user-preferences-button" class="user-preferences-button button" title="Modify user preferences"><span id="avatarbutton" class="avatarbutton name">You!</span><div id="preferences_menu_holder" class="preferences_menu_holder hidden"><ol id="preferences_menu" class="preferences_menu" style="font-family: 'Roboto Serif', serif;">
<li data-env="avatar" tabindex="-1">Choose avatar<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden avatar">
<li data-val="You!" tabindex="-1">
<span id="theYou!" class="avatarcheck">✔️</span>You!</li>
<li data-val="😺" tabindex="-1">
<span id="the😺" class="avatarcheck"></span>😺</li>
<li data-val="👤" tabindex="-1">
<span id="the👤" class="avatarcheck"></span>👤</li>
<li data-val="👽" tabindex="-1">
<span id="the👽" class="avatarcheck"></span>👽</li>
<li data-val="🐶" tabindex="-1">
<span id="the🐶" class="avatarcheck"></span>🐶</li>
<li data-val="🐼" tabindex="-1">
<span id="the🐼" class="avatarcheck"></span>🐼</li>
<li data-val="🌈" tabindex="-1">
<span id="the🌈" class="avatarcheck"></span>🌈</li>
</ol>
</li>
<li data-env="fontfamily" tabindex="-1">Font family<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden fontfamily">
<li data-val="face" data-change="OS" tabindex="-1" style="font-family: 'Open Sans'">
<span id="theOS" class="ffcheck">✔️</span><span class="name">Open Sans</span><span class="sample">AaBbCc 123 PreTeXt</span>
</li>
<li data-val="face" data-change="RS" tabindex="-1" style="font-family: 'Roboto Serif'">
<span id="theRS" class="ffcheck"></span><span class="name">Roboto Serif</span><span class="sample">AaBbCc 123 PreTeXt</span>
</li>
</ol>
</li>
<li data-env="font" tabindex="-1">Adjust font<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden fonts">
<li>Size</li>
<li><span id="thesize">12</span></li>
<li data-val="size" data-change="-1" tabindex="-1" style="font-size: 80%">Smaller</li>
<li data-val="size" data-change="1" tabindex="-1" style="font-size: 110%">Larger</li>
<li>Width</li>
<li><span id="thewdth">100</span></li>
<li data-val="wdth" data-change="-5" tabindex="-1" style="font-variation-settings: 'wdth' 60">narrower</li>
<li data-val="wdth" data-change="5" tabindex="-1" style="font-variation-settings: 'wdth' 150">wider</li>
<li>Weight</li>
<li><span id="thewght">400</span></li>
<li data-val="wght" data-change="-50" tabindex="-1" style="font-weight: 200">thinner</li>
<li data-val="wght" data-change="50" tabindex="-1" style="font-weight: 700">heavier</li>
<li>Letter spacing</li>
<li>
<span id="thelspace">0</span><span class="byunits">/200</span>
</li>
<li data-val="lspace" data-change="-1" tabindex="-1">closer</li>
<li data-val="lspace" data-change="1" tabindex="-1">f a r t h e r</li>
<li>Word spacing</li>
<li>
<span id="thewspace">0</span><span class="byunits">/50</span>
</li>
<li data-val="wspace" data-change="-1" tabindex="-1">smaller gap </li>
<li data-val="wspace" data-change="1" tabindex="-1">larger gap</li>
<li>Line Spacing</li>
<li>
<span id="theheight">135</span><span class="byunits">/100</span>
</li>
<li data-val="height" data-change="-5" tabindex="-1" style="line-height: 1">closer<br>together</li>
<li data-val="height" data-change="5" tabindex="-1" style="line-height: 1.75">further<br>apart</li>
</ol>
</li>
<li data-env="atmosphere" tabindex="-1">Light/dark mode<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden atmosphere">
<li data-val="default" tabindex="-1">
<span id="thedefault" class="atmospherecheck">✔️</span>default</li>
<li data-val="pastel" tabindex="-1">
<span id="thepastel" class="atmospherecheck"></span>pastel</li>
<li data-val="darktwilight" tabindex="-1">
<span id="thedarktwilight" class="atmospherecheck"></span>twilight</li>
<li data-val="dark" tabindex="-1">
<span id="thedark" class="atmospherecheck"></span>dark</li>
<li data-val="darkmidnight" tabindex="-1">
<span id="thedarkmidnight" class="atmospherecheck"></span>midnight</li>
</ol>
</li>
<li data-env="ruler" tabindex="-1">Reading ruler<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden ruler">
<li data-val="none" tabindex="-1">
<span id="thenone" class="rulercheck">✔️</span>none</li>
<li data-val="underline" tabindex="-1">
<span id="theunderline" class="rulercheck"></span>underline</li>
<li data-val="lunderline" tabindex="-1">
<span id="thelunderline" class="rulercheck"></span>L-underline</li>
<li data-val="greybar" tabindex="-1">
<span id="thegreybar" class="rulercheck"></span>grey bar</li>
<li data-val="lightbox" tabindex="-1">
<span id="thelightbox" class="rulercheck"></span>light box</li>
<li data-val="sunrise" tabindex="-1">
<span id="thesunrise" class="rulercheck"></span>sunrise</li>
<li data-val="sunriseunderline" tabindex="-1">
<span id="thesunriseunderline" class="rulercheck"></span>sunrise underline</li>
<li class="moveQ">Motion by:</li>
<li data-val="mouse" tabindex="-1">
<span id="themouse" class="motioncheck">✔️</span>follow the mouse</li>
<li data-val="arrow" tabindex="-1">
<span id="thearrow" class="motioncheck"></span>up/down arrows - not yet</li>
<li data-val="eye" tabindex="-1">
<span id="theeye" class="motioncheck"></span>eye tracking - not yet</li>
</ol>
</li>
</ol></div></button><span class="treebuttons"><a class="previous-button button" href="peer_instruction.html" title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="up-button button" href="InstructorGuide.html" title="Up"><span class="icon">^</span><span class="name">Up</span></a><a class="next-button button" href="backmatter-1.html" title="Next"><span class="name">Next</span><span class="icon">&gt;</span></a></span><div class="searchbox"><div class="searchwidget">
<input id="ptxsearch" class="ptxsearch" type="text" name="terms" placeholder="Search" onchange="doSearch()"><button id="searchbutton" class="searchbutton" type="button" onclick="doSearch()">🔍</button>
</div></div></nav><div id="latex-macros" class="hidden-content process-math" style="display:none"><span class="process-math">\(
\newcommand{\lt}{&lt;}
\newcommand{\gt}{&gt;}
\newcommand{\amp}{&amp;}
\definecolor{fillinmathshade}{gray}{0.9}
\newcommand{\fillinmath}[1]{\mathchoice{\colorbox{fillinmathshade}{$\displaystyle     \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\textstyle        \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptstyle      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptscriptstyle\phantom{\,#1\,}$}}}
\)</span></div>
<div class="ptx-page">
<div id="ptx-sidebar" class="ptx-sidebar"><nav id="ptx-toc" class="ptx-toc depth4"><ul class="structural">
<li><div class="toc-item"><a href="frontmatter-1.html" class="internal"><span class="title">Front Matter</span></a></div></li>
<li>
<div class="toc-item"><a href="Introduction.html" class="internal"><span class="codenumber">1</span> <span class="title">Introduction</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="Introduction.html#prose-philosophy" class="internal"><span class="codenumber">1.1</span> <span class="title">The Runestone Philosophy</span></a></div></li>
<li><div class="toc-item"><a href="Introduction.html#section-2" class="internal"><span class="codenumber">1.2</span> <span class="title">Courses Versus Books</span></a></div></li>
<li><div class="toc-item"><a href="Introduction.html#section-3" class="internal"><span class="codenumber">1.3</span> <span class="title">Privacy</span></a></div></li>
<li><div class="toc-item"><a href="Introduction.html#section-4" class="internal"><span class="codenumber">1.4</span> <span class="title">Why am I seeing Ads</span></a></div></li>
<li><div class="toc-item"><a href="Introduction.html#a-brief-history" class="internal"><span class="codenumber">1.5</span> <span class="title">A Brief History of Runestone</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="registration_register-for-this-course.html" class="internal"><span class="codenumber">2</span> <span class="title">Register for a Course</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="registration_register-for-this-course.html#section-6" class="internal"><span class="codenumber">2.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="registration_register-for-this-course.html#registration_first-time-registration" class="internal"><span class="codenumber">2.2</span> <span class="title">First-time Registration</span></a></div></li>
<li><div class="toc-item"><a href="registration_register-for-this-course.html#registration_add-an-additional-course" class="internal"><span class="codenumber">2.3</span> <span class="title">Add an Additional Course</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-3.html" class="internal"><span class="codenumber">3</span> <span class="title">Creating a Course</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-3.html#custom-courses_what-is-a-custom-course" class="internal"><span class="codenumber">3.1</span> <span class="title">What is a course?</span></a></div></li>
<li><div class="toc-item"><a href="chapter-3.html#custom-courses_how-to-create-a-custom-course" class="internal"><span class="codenumber">3.2</span> <span class="title">How to Create a course</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-4.html" class="internal"><span class="codenumber">4</span> <span class="title">Managing Your Course</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="chapter-4.html#registration_adding-students-to-your-course" class="internal"><span class="codenumber">4.1</span> <span class="title">Managing Students in your Course</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-4.html#subsection-1" class="internal"><span class="codenumber">4.1.1</span> <span class="title">Student Self Registration</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#subsection-2" class="internal"><span class="codenumber">4.1.2</span> <span class="title">Pre-register your students</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#subsection-3" class="internal"><span class="codenumber">4.1.3</span> <span class="title">Removing Students</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#subsection-4" class="internal"><span class="codenumber">4.1.4</span> <span class="title">Resetting Passwords</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-4.html#registration_privacy-and-runestone" class="internal"><span class="codenumber">4.2</span> <span class="title">Privacy and Runestone</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-4.html#registration_setting-up-a-class-with-no-pii" class="internal"><span class="codenumber">4.2.1</span> <span class="title">Setting up a class with no PII</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#registration_what-other-data-does-runestone-collect" class="internal"><span class="codenumber">4.2.2</span> <span class="title">What other Data Does Runestone Collect</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-4.html#lti_integration" class="internal"><span class="codenumber">4.3</span> <span class="title">Integrating with your LMS</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-4.html#subsection-7" class="internal"><span class="codenumber">4.3.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#subsection-8" class="internal"><span class="codenumber">4.3.2</span> <span class="title">Phase 1 (All LMS platforms)</span></a></div></li>
<li>
<div class="toc-item"><a href="chapter-4.html#subsection-9" class="internal"><span class="codenumber">4.3.3</span> <span class="title">Phase 2 (LMS-specific Steps)</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-4.html#subsubsection-1" class="internal"><span class="title">Canvas</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#subsubsection-2" class="internal"><span class="title">Moodle</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li><div class="toc-item"><a href="chapter-4.html#add_instructor" class="internal"><span class="codenumber">4.4</span> <span class="title">Add and Instructor or TA</span></a></div></li>
<li><div class="toc-item"><a href="chapter-4.html#course_settings" class="internal"><span class="codenumber">4.5</span> <span class="title">Course Settings</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-5.html" class="internal"><span class="codenumber">5</span> <span class="title">Monitoring your Student’s Learning</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-5.html#instructor-interface_the-instructor-interface" class="internal"><span class="codenumber">5.1</span> <span class="title">The Instructor Interface</span></a></div></li>
<li>
<div class="toc-item"><a href="chapter-5.html#instructor-interface_view-student-activity" class="internal"><span class="codenumber">5.2</span> <span class="title">View Student Activity</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-5.html#activity_intro" class="internal"><span class="codenumber">5.2.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="chapter-5.html#instructor-interface_chapter-progress" class="internal"><span class="codenumber">5.2.2</span> <span class="title">Chapter Progress</span></a></div></li>
<li><div class="toc-item"><a href="chapter-5.html#instructor-interface_student-progress-tab" class="internal"><span class="codenumber">5.2.3</span> <span class="title">Student Progress Tab</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="working-with-assignments-chap.html" class="internal"><span class="codenumber">6</span> <span class="title">Working with Assignments</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_assignments" class="internal"><span class="codenumber">6.1</span> <span class="title">Understanding Assignments</span></a></div></li>
<li>
<div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_create-an-assignment" class="internal"><span class="codenumber">6.2</span> <span class="title">Create an Assignment</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="working-with-assignments-chap.html#create_assignment_intro" class="internal"><span class="codenumber">6.2.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_creating-a-reading-assignment" class="internal"><span class="codenumber">6.2.2</span> <span class="title">Creating a Reading Assignment</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_creating-a-problem-assignment" class="internal"><span class="codenumber">6.2.3</span> <span class="title">Creating a Problem Assignment</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_copying-assignments" class="internal"><span class="codenumber">6.3</span> <span class="title">Copying Assignments</span></a></div></li>
<li>
<div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_exams-and-quizzes-in-runestone" class="internal"><span class="codenumber">6.4</span> <span class="title">Exams and Quizzes in Runestone</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="working-with-assignments-chap.html#create_quiz_intro" class="internal"><span class="codenumber">6.4.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_resetting-an-exam" class="internal"><span class="codenumber">6.4.2</span> <span class="title">Resetting an Exam</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_grading" class="internal"><span class="codenumber">6.5</span> <span class="title">Grading</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="working-with-assignments-chap.html#grading_intro" class="internal"><span class="codenumber">6.5.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_autograde" class="internal"><span class="codenumber">6.5.2</span> <span class="title">Autograde</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_calculate-totals" class="internal"><span class="codenumber">6.5.3</span> <span class="title">Calculate Totals</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_release-grades" class="internal"><span class="codenumber">6.5.4</span> <span class="title">Release Grades</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_how-a-student-sees-grades-and-comments" class="internal"><span class="codenumber">6.5.5</span> <span class="title">How a Student Sees Grades and Comments</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_the-practice-interface" class="internal"><span class="codenumber">6.6</span> <span class="title">The Practice Interface</span></a></div></li>
<li><div class="toc-item"><a href="working-with-assignments-chap.html#instructor-interface_exercises" class="internal"><span class="codenumber">6.7</span> <span class="title">Exercises</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="peer_instruction.html" class="internal"><span class="codenumber">7</span> <span class="title">Peer Instruction</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="peer_instruction.html#section-24" class="internal"><span class="codenumber">7.1</span> <span class="title">Peer instruction in Runestone</span></a></div></li>
<li><div class="toc-item"><a href="peer_instruction.html#section-25" class="internal"><span class="codenumber">7.2</span> <span class="title">Creating Peer instruction Assignments</span></a></div></li>
<li><div class="toc-item"><a href="peer_instruction.html#section-26" class="internal"><span class="codenumber">7.3</span> <span class="title">Virtual Peer Instruction Workshop</span></a></div></li>
</ul>
</li>
<li class="active">
<div class="toc-item"><a href="chapter-8.html" class="internal"><span class="codenumber">8</span> <span class="title">Writing your own Exercises</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="chapter-8.html#writing-exercises_adding-an-exercise-to-the-exercise-bank" class="internal"><span class="codenumber">8.1</span> <span class="title">Adding an Exercise to the Exercise Bank</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-8.html#writing_exer_intro" class="internal"><span class="codenumber">8.1.1</span> <span class="title">Introduction</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_unit-testing-exercises" class="internal"><span class="codenumber">8.1.2</span> <span class="title">Unit Testing exercises</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_advanced-activecode-options" class="internal"><span class="codenumber">8.1.3</span> <span class="title">Advanced activecode Options</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_for-languages-outside-the-browser" class="internal"><span class="codenumber">8.1.4</span> <span class="title">For languages outside the browser</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_multiple-choice-questions" class="internal"><span class="codenumber">8.2</span> <span class="title">Multiple Choice Questions</span></a></div></li>
<li>
<div class="toc-item"><a href="chapter-8.html#writing-exercises_mixed-up-code-parsons" class="internal"><span class="codenumber">8.3</span> <span class="title">Mixed-Up Code (Parsons)</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_parsons-options" class="internal"><span class="codenumber">8.3.1</span> <span class="title">Parsons Options</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_parsons-distractor-types" class="internal"><span class="codenumber">8.3.2</span> <span class="title">Parsons Distractor Types</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_example-with-paired-distractors" class="internal"><span class="codenumber">8.3.3</span> <span class="title">Example with Paired Distractors</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_example-with-unpaired-distractor" class="internal"><span class="codenumber">8.3.4</span> <span class="title">Example with Unpaired Distractor</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_pro-tip-test-your-exercises-locally" class="internal"><span class="codenumber">8.4</span> <span class="title">Pro Tip – Test Your Exercises Locally</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_add-your-own-question" class="internal"><span class="codenumber">8.5</span> <span class="title">Add your own Question</span></a></div></li>
<li><div class="toc-item"><a href="chapter-8.html#writing-exercises_exercises" class="internal"><span class="codenumber">8.6</span> <span class="title">Exercises</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="backmatter-1.html" class="internal"><span class="title">Back Matter</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="index-1.html" class="internal"><span class="title">Index</span></a></div></li>
<li>
<div class="toc-item"><a href="welcome.html" class="internal"><span class="codenumber">A</span> <span class="title">Welcome to the Runestone Community</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="welcome.html#help-and-support" class="internal"><span class="codenumber">A.1</span> <span class="title">Help and Support</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#drop-in" class="internal"><span class="codenumber">A.2</span> <span class="title">Weekly Drop-Ins</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#github-issues" class="internal"><span class="codenumber">A.3</span> <span class="title">Feature Requests and Reporting Problems</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#contributing" class="internal"><span class="codenumber">A.4</span> <span class="title">Contributing</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="welcome.html" class="internal"><span class="codenumber">B</span> <span class="title">Frequently Asked Questions</span></a></div>
<ul class="structural"><li>
<div class="toc-item"><a href="welcome.html#faq" class="internal"><span class="codenumber">B.1</span> <span class="title">Frequently Asked Questions</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="welcome.html#_why-can-t-i-access-the-instructor-page" class="internal"><span class="codenumber">B.1.1</span> <span class="title">Why Can’t I access the Instructor Page?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-become-an-instructor" class="internal"><span class="codenumber">B.1.2</span> <span class="title">How do I become an instructor?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-learn-more-about-the-instructor-interface" class="internal"><span class="codenumber">B.1.3</span> <span class="title">How do I learn more about the instructor interface</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_help-i-cannot-log-in-every-time-i-try-to-log-in-i-have-to-reset-my-password" class="internal"><span class="codenumber">B.1.4</span> <span class="title">Help! I cannot log in. Every time I try to log in, I have to reset my password!</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-m-not-getting-the-password-reset-email-can-you-reset-my-password" class="internal"><span class="codenumber">B.1.5</span> <span class="title">I’m not getting the password reset email, can you reset my password?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-can-t-log-in-to-my-coursera-course-what-do-i-do" class="internal"><span class="codenumber">B.1.6</span> <span class="title">I can’t log in to my Coursera course, what do I do?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-add-students-to-my-course" class="internal"><span class="codenumber">B.1.7</span> <span class="title">How do I add students to my course?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_what-do-those-green-checkmarks-and-orange-circles-mean" class="internal"><span class="codenumber">B.1.8</span> <span class="title">What do those green checkmarks and orange circles mean?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-come-there-is-no-green-checkmark-for-the-exercises-pages" class="internal"><span class="codenumber">B.1.9</span> <span class="title">How come there is no Green Checkmark for the Exercises pages?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_where-do-i-find-answers-to-all-the-homework-questions" class="internal"><span class="codenumber">B.1.10</span> <span class="title">Where do I find answers to all the homework questions?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_can-i-build-my-own-course-and-host-it-here" class="internal"><span class="codenumber">B.1.11</span> <span class="title">Can I build my own course and host it here?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-build-my-own-course" class="internal"><span class="codenumber">B.1.12</span> <span class="title">How do I build my own course?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_is-this-site-reliable-enough-to-use-in-class" class="internal"><span class="codenumber">B.1.13</span> <span class="title">Is this site reliable enough to use in class?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_why-did-you-change-domain-names" class="internal"><span class="codenumber">B.1.14</span> <span class="title">Why did you change domain names?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-want-to-reuse-my-course-from-last-year-what-should-i-do" class="internal"><span class="codenumber">B.1.15</span> <span class="title">I want to reuse my course from last year, what should I do?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-update-my-course-to-get-the-latest-bug-fixes" class="internal"><span class="codenumber">B.1.16</span> <span class="title">How do I update my course to get the latest bug fixes?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-was-just-experimenting-and-i-want-to-delete-my-course" class="internal"><span class="codenumber">B.1.17</span> <span class="title">I was just experimenting and I want to delete my course</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_what-if-i-want-to-add-a-new-section-or-chapter" class="internal"><span class="codenumber">B.1.18</span> <span class="title">What if I want to add a new section or chapter?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_what-if-i-want-to-add-my-own-exercises" class="internal"><span class="codenumber">B.1.19</span> <span class="title">What if I want to add my own exercises?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_my-students-are-completing-work-but-it-is-not-showing-up-help" class="internal"><span class="codenumber">B.1.20</span> <span class="title">My students are completing work but it is not showing up. Help!</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_how-do-i-change-courses-or-register-for-an-additional-course" class="internal"><span class="codenumber">B.1.21</span> <span class="title">How do I change courses or register for an additional course?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_what-course-name-should-i-use-when-i-register" class="internal"><span class="codenumber">B.1.22</span> <span class="title">What course name should I use when I register?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_what-version-of-python-does-your-book-use" class="internal"><span class="codenumber">B.1.23</span> <span class="title">What version of Python does your book use?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-think-there-is-a-bug-in-your-book-what-should-i-do" class="internal"><span class="codenumber">B.1.24</span> <span class="title">I think there is a bug in your book what should I do?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_i-have-a-question-that-is-not-covered-here" class="internal"><span class="codenumber">B.1.25</span> <span class="title">I have a question that is not covered here!</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_my-internet-is-slow-i-want-to-use-runestone-offline" class="internal"><span class="codenumber">B.1.26</span> <span class="title">My Internet is slow, I want to use Runestone offline</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#_why-can-t-i-just-get-a-pdf" class="internal"><span class="codenumber">B.1.27</span> <span class="title">Why can’t I just get a PDF?</span></a></div></li>
<li><div class="toc-item"><a href="welcome.html#timezone" class="internal"><span class="codenumber">B.1.28</span> <span class="title">How does Runestone handle timezones?</span></a></div></li>
</ul>
</li></ul>
</li>
</ul>
</li>
</ul></nav></div>
<main class="ptx-main"><div id="ptx-content" class="ptx-content">
<section class="chapter" id="chapter-8"><h2 class="heading">
<span class="type">Chapter</span> <span class="codenumber">8</span> <span class="title">Writing your own Exercises</span>
</h2>
<section class="section" id="writing-exercises_adding-an-exercise-to-the-exercise-bank"><h3 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">8.1</span> <span class="title">Adding an Exercise to the Exercise Bank</span>
</h3>
<section class="subsection" id="writing_exer_intro"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.1.1</span> <span class="title">Introduction</span>
</h4>
<div class="video-box" style="width: 100%;padding-top: 56.25%; margin-left: 0%; margin-right: 0%;"><iframe id="l84cb4wtM8g" class="video" allowfullscreen="" src="https://www.youtube-nocookie.com/embed/l84cb4wtM8g?&amp;modestbranding=1&amp;rel=0"></iframe></div>
<div class="para" id="p-226">Although all of the Runestone books have a pretty good selection of exercises, you can never have too many exercises. We hope to crowd source an enormous collection of questions and problems in our Exercise database. You can help by adding custom exercises for your own course to the database! Once your exercise is in the database it is available for others to find using the search feature on the create assignment page.</div>
<div class="para" id="p-227">Since the most common exercises are programming exercises let’s look at the activecode directive in detail. Here is the full directive with every possible option. This is what you would see if you add an activecode exercise from the web interface. Let’s not get bogged down in the details here, let’s look at a much simpler example.</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. activecode:: uniqueid 'nocanvas': directives.flag,
:nopre: do not create an output component
:above: put the canvas above the code
:autorun: run this activecode as soon as the page is loaded
:autograde: typically unittest
:caption: caption under the active code
:include: invisibly include code from another activecode
:hidecode: Don:t show the editor initially
:language: python, html, javascript, java, python2, python3
:tour_1: audio tour track
:tour_2: audio tour track
:tour_3: audio tour track
:tour_4: audio tour track
:tour_5: audio tour track
:nocodelens: Do not show the codelens button
:coach: Show the codecoach button
:timelimit: set the time limit for this program
:stdin: : A file to simulate stdin (java, python2, python3)
:datafile: : A datafile for the program to read (java, python2, python3)
:sourcefile: : source files (java, python2, python3)
:available_files: : other additional files (java, python2, python3)

If this is a homework problem instead of an example in the text
then the assignment text should go here. The assignment text ends with
the line containing four tilde ~
~~~~
print("hello world")
====
print("Hidden code, such as unit tests come after the four = signs")
</code></pre></div>
<div class="para" id="p-228">Let’s make an exercise to have the student compute the sum of the first N numbers.</div>
<article class="exercise exercise-like" id="sigcse_ex1"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.1.1</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="sigcse_ex1">
<div class="ac_question col-md-12" id="sigcse_ex1_question"><div class="para" id="p-229">Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="sigcse_ex1_editor" data-question_label="" data-codelens="true">def sum_first_n(N):
# your code here
</textarea>
</div></div></div></article><div class="para" id="p-230">The source code for the above looks like this:</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. activecode:: sigcse_ex1
:language: python

Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3
~~~~
def sum_first_n(N):
# your code here
</code></pre></div>
<div class="para" id="p-231">Indentation and whitespace is important. In the previous example, you will see that the optional directive <code class="code-inline tex2jax_ignore">:language:</code> is indented by four spaces. from the previous line. It must be indented at least three spaces to line up with the “a” in activecode, but using 4 keeps it consistent with my own personal python indentation style. Everything else in the body of the directive must also be indented to match the indentation of the optional parameters.</div>
<div class="para" id="p-232">You can enter this into a textbox by clicking on the “Write” button under the Problems area.</div>
<div class="image-box" style="width: 50%; margin-left: 25%; margin-right: 25%;"><img src="external/Figures/enter_rst.png" class="contained"></div>
<div class="para" id="p-233">After you enter the text you need to compile and render the text. You can get a preview of your question on the right of the page.</div>
<div class="para" id="p-234">The first part of the directive body are the instructions for the student. You can use any valid restructured text (rst) in this part of the directive including embedded images. You separate the instructions from any code you want to give the student with <code class="code-inline tex2jax_ignore">~~~~</code> thats four tildes. The next part of the body is any code you want to provide. Just put in your Python or other language as you would normally write it.</div>
<div class="para" id="p-235">To keep an exercise page looking a bit neater you can substitute the <code class="code-inline tex2jax_ignore">.. actex::</code> directive for the <code class="code-inline tex2jax_ignore">.. activecode::</code> directive. Making that substitution gives us the following:</div>
<article class="exercise exercise-like" id="sigcse_ex2"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.1.2</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="sigcse_ex2">
<div class="ac_question col-md-12" id="sigcse_ex2_question"><div class="para" id="p-236">Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="sigcse_ex2_editor" data-question_label="" data-codelens="true">def sum_first_n(N):
# your code here
</textarea>
</div></div></div></article></section><section class="subsection" id="writing-exercises_unit-testing-exercises"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.1.2</span> <span class="title">Unit Testing exercises</span>
</h4>
<div class="para" id="p-237">Of course, as you have learned in the previous section, the real killer feature of these exercises is to be able to create your own unit tests, and have the grader autograde the assignment.</div>
<div class="para" id="p-238">Let’s expand our example to include some simple unittests. We can do this by adding a hidden block of code to our previous example that uses the standard Python unittest framework.</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-py">from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(add(2,2),4,"A feedback string when the test fails")
self.assertAlmostEqual(add(2.0,3.0), 5.0, 1, "Try adding your parmeters")

myTests().main()
</code></pre></div>
<div class="para" id="p-239">If you are not familiar with Python unittests they are pretty easy to write. You create your own class that is a subclass of TestCase, or in our work TestCaseGui so we get some graphical output. Your tests are all methods of the class and must start with the word “test”. There are a host of assertXXXX functions that you can use. Check out the <a class="external" href="https://docs.python.org/2/library/unittest.html#assert-methods" target="_blank">unittest documentation here</a><a href="" data-knowl="" class="id-ref fn-knowl original" data-refid="hk-fn-9" id="fn-9"><sup> 1 </sup></a>
</div>
<div class="para" id="p-240">The second important addition in this example is the :autograde: option. This will allow the grade to make use of your unittests and assign a grade automatically for this particular problem.</div>
<article class="exercise exercise-like" id="sigcse_ex3"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.1.3</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="sigcse_ex3">
<div class="ac_question col-md-12" id="sigcse_ex3_question"><div class="para" id="p-241">Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="sigcse_ex3_editor" data-question_label="" data-codelens="true">def sum_first_n(N):
# your code here
====
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(sum_first_n(4),6,feedback="0 + 1 + 2 + 3 == 6")
self.assertEqual(sum_first_n(0),0,feedback="summing 0 numbers should be 0")

myTests().main()
</textarea>
</div></div></div></article><div class="para" id="p-242">Let’s try to add another test to the example above. This time we’ll show the unittests in the active code window to make it easy for us to test. I strongly recommend you do this in a scratch activecode window when you are writing a question. It will save lots of frustrating wait time as you work out the details of your testing.</div>
<div class="para" id="p-243">Let’s try to add another test to the example above. This time we’ll show the unittests in the active code window to make it easy for us to test. I strongly recommend you do this in a scratch activecode window when you are writing a question. It will save lots of frustrating wait time as you work out the details of your testing.</div>
<article class="exercise exercise-like" id="sigcse_ex4"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.1.4</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="sigcse_ex4">
<div class="ac_question col-md-12" id="sigcse_ex4_question"><div class="para" id="p-244">Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="sigcse_ex4_editor" data-question_label="" data-codelens="true">def sum_first_n(N):
# your code here

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(sum_first_n(4),6,feedback="0 + 1 + 2 + 3 == 6")
self.assertEqual(sum_first_n(0),0,feedback="summing 0 numbers should be 0")

myTests().main()
</textarea>
</div></div></div></article></section><section class="subsection" id="writing-exercises_advanced-activecode-options"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.1.3</span> <span class="title">Advanced activecode Options</span>
</h4>
<div class="para logical" id="p-245"><ul class="disc">
<li id="li-121"><div class="para" id="p-derived-li-121">include <div class="para" id="p-246">– This option lets you include other activecodes in the current</div>
</div></li>
<li id="li-122"><div class="para" id="p-derived-li-122">timelimit <div class="para" id="p-247">seconds – What to do when students create an infinite loop and lock up their browser? Just wait a bit, every run of Python has a built in time limit of 30 seconds. Some things might take longer than this, so if you know an example or assignment is going to take longer, then you can set a higher time limit with this option.</div>
</div></li>
<li id="li-123"><div class="para" id="p-derived-li-123">nocodelens <div class="para" id="p-248">– codelens is an awesome addition, but it does not work with very many libraries. This is part of the sandboxed security of the codelens server. The most common one to be aware of is the turtle module. If you are doing a turtle example or assigning a turtle problem then you should set this flag so the “Show Codelens” button is hidden.</div>
</div></li>
</ul></div></section><section class="subsection" id="writing-exercises_for-languages-outside-the-browser"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.1.4</span> <span class="title">For languages outside the browser</span>
</h4>
<div class="para logical" id="p-249"><ul class="disc"><li id="li-124"><div class="para" id="p-derived-li-124">language <div class="para" id="p-250">– As you saw earlier, Runestone supports Java, python2, and python3 in a sandboxed server environment environment.</div>
</div></li></ul></div>
<div class="para" id="p-251">If you choose any of the above, the code from the window is packaged up and set to a separate server for compilation and execution. There are a few options to activecode that only apply to these languages.</div>
<div class="para logical" id="p-252"><ul class="disc">
<li id="li-125"><div class="para" id="p-derived-li-125">datafile <div class="para" id="p-253">You can provide an id to a datafile that will be sent along with your Java code</div>
</div></li>
<li id="li-126"><div class="para" id="p-derived-li-126">sourcefile <div class="para" id="p-254">You can provide additional source files that should be compiled along with the java file you upload.</div>
</div></li>
<li id="li-127"><div class="para" id="p-derived-li-127">available_files <div class="para" id="p-255">You can provide additional binary files to link into the final executable.</div>
</div></li>
</ul></div></section></section><section class="section" id="writing-exercises_multiple-choice-questions"><h3 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">8.2</span> <span class="title">Multiple Choice Questions</span>
</h3>
<div class="para" id="p-256">Another common question type is the multiple choice question.</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. mchoice:: uniqueid
    :multiple_answers: boolean  [optional]
    :random: boolean [optional]
    :answer_a: possible answer  -- what follows _ is label
    :answer_b: possible answer
    :answer_c: possible answer
    :answer_d: possible answer
    :answer_e: possible answer
    :correct: letter of correct answer or list of correct answer letters (in case of multiple answe
    rs)
    :feedback_a: displayed if a is picked
    :feedback_b: displayed if b is picked
    :feedback_c: displayed if c is picked
    :feedback_d: displayed if d is picked
    :feedback_e: displayed if e is picked

    Question text   ...
</code></pre></div>
<div class="para" id="p-257">Here is an example from a Data Structures class:</div>
<article class="exercise exercise-like" id="stack_2"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.2.1</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone"><ul data-component="multiplechoice" id="stack_2" data-multipleanswers="false">
<div class="para" id="p-258">3-1-1: Given the following sequence of stack operations, what is the top item on the stack when the sequence is complete?</div> <div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-py">m = Stack()
m.push('x')
m.push('y')
m.push('z')
while not m.isEmpty():
   m.pop()
   m.pop()
</code></pre></div>
<li data-component="answer" id="stack_2_opt_a"><div class="para" id="p-259">’x’</div></li>
<li data-component="feedback" id="stack_2_opt_a"><div class="para" id="p-260">You may want to check out the docs for isEmpty</div></li>
<li data-component="answer" id="stack_2_opt_b"><div class="para" id="p-261">the stack is empty</div></li>
<li data-component="feedback" id="stack_2_opt_b"><div class="para" id="p-262">There is an odd number of things on the stack but each time through the loop 2 things are popped.</div></li>
<li data-component="answer" id="stack_2_opt_c" data-correct=""><div class="para" id="p-263">an error will occur</div></li>
<li data-component="feedback" id="stack_2_opt_c"><div class="para" id="p-264">Good Job.</div></li>
<li data-component="answer" id="stack_2_opt_d"><div class="para" id="p-265">’z’</div></li>
<li data-component="feedback" id="stack_2_opt_d"><div class="para" id="p-266">You may want to check out the docs for isEmpty</div></li>
</ul></div></div></article><div class="para" id="p-267">The source code is here:</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. mchoice:: stack_2
    :answer_a: 'x'
    :answer_b: the stack is empty
    :answer_c: an error will occur
    :answer_d: 'z'
    :correct: c
    :feedback_a: You may want to check out the docs for isEmpty
    :feedback_b: There is an odd number of things on the stack but each time through the loop 2 things are popped.
    :feedback_c: Good Job.
    :feedback_d: You may want to check out the docs for isEmpty

    Given the following sequence of stack operations, what is the top item on the stack when the sequence is complete?

    .. code-block:: python

        m = Stack()
        m.push('x')
        m.push('y')
        m.push('z')
        while not m.isEmpty():
            m.pop()
            m.pop()
</code></pre></div></section><section class="section" id="writing-exercises_mixed-up-code-parsons"><h3 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">8.3</span> <span class="title">Mixed-Up Code (Parsons)</span>
</h3>
<section class="introduction" id="introduction-5"><div class="para" id="p-268">Mixed-up code or Parsons problems provide the correct code to solve a problem, but the code is broken into blocks and mixed up.</div> <div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">   .. parsonsprob:: unqiue_problem_id_here
      :maxdist:
      :order:
      :language:
      :noindent:
      :adaptive:
      :numbered: left

      Instructions for the user. These can include a textual description of how to solve the problem. You must leave a blank line before this.
      -----
      def findmax(alist):
      =====
      if len(alist) == 0:
      return None
      =====
      curmax = alist[0]
      for item in alist:
      =====
      if item &amp;gt; curmax:
      =====
      curmax = item
      =====
      return curmax
</code></pre></div> <div class="para" id="p-269">Create a working program and then paste the code for it into the editor.  Indent the code so that the left edge lines up with the options. Indent each line with 4 spaces beyond the previous line. Separate the blocks with “=====” which must line up under the options. Put the instructions before the code after a blank line and then followed by “----”.</div></section><section class="subsection" id="writing-exercises_parsons-options"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.3.1</span> <span class="title">Parsons Options</span>
</h4>
<div class="para" id="p-270">Options are indented under the Parsons directive (under the p).</div>
<figure class="table table-like" id="writing-exercises_id1"><figcaption><span class="type">Table</span><span class="space"> </span><span class="codenumber">8.3.1<span class="period">.</span></span><span class="space"> </span></figcaption><div class="tabular-box natural-width"><table class="tabular">
<tr class="header-horizontal">
<th scope="col" class="l m b0 r0 l0 t0 lines">Option</th>
<th scope="col" class="l m b0 r0 l0 t0 lines">Description</th>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">maxdist</td>
<td class="l m b0 r0 l0 t0 lines">The maximum number of distractors to use. They will be picked at random</td>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">order</td>
<td class="l m b0 r0 l0 t0 lines">The order for the lines, they are displayed in a random order normally</td>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">language</td>
<td class="l m b0 r0 l0 t0 lines">Specify the language: java, python</td>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">noindent</td>
<td class="l m b0 r0 l0 t0 lines">Provide the indentation for the user by adding spaces to the left of the code</td>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">adaptive</td>
<td class="l m b0 r0 l0 t0 lines">Provide help is the user is struggling and modify the difficulty of the problem based on the user’s performance on the previous problem</td>
</tr>
<tr>
<td class="l m b0 r0 l0 t0 lines">numbered</td>
<td class="l m b0 r0 l0 t0 lines">Show numbered labels to the left of the code if you add left or to the right when you add right</td>
</tr>
</table></div></figure></section><section class="subsection" id="writing-exercises_parsons-distractor-types"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.3.2</span> <span class="title">Parsons Distractor Types</span>
</h4>
<div class="para" id="p-271">You can include distractor blocks in the problem. A distractor is code that isn’t needed in a correct solution, such as code with a syntax error. Add a distractor block after the correct code block. Distractors can either be paired or unpaired. For paired distractors use #paired at end of the first line of code in the distractor block. For unpaired distractors add #distractor.</div></section><section class="subsection" id="writing-exercises_example-with-paired-distractors"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.3.3</span> <span class="title">Example with Paired Distractors</span>
</h4>
<div class="para" id="p-272">Here is an example with paired distractors from a data oriented intermediate programming course in Python.</div>
<article class="exercise exercise-like" id="mt1dict1ex"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.3.2</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone parsons_section" style="max-width: none;"><div data-component="parsons" class="parsons" id="mt1dict1ex">
<div class="parsons_question parsons-text"><div class="para" id="p-273">Complete the function greater_dictionary. Given a dictionary d and an integer cutoff, return a dictionary that contains only the key-value pairs where they key is greater than or equal to cutoff.</div></div>
<pre class="parsonsblocks" data-question_label="" style="visibility: hidden;" data-language="python" data-noindent="true">def greater_dictionary(d, cutoff):
---
def greater_dictionary(self, d, cutoff):  #paired
---
 result = {}
---
 for key in d.keys():
---
 for key in range(d):  #paired
---
 if key &gt;= cutoff:
---
 if key &gt; cutoff:  #paired
---
 result[key] = d[key]
---
 d[key] = result[key]  #paired
---
 return result</pre>
</div></div></div></article><div class="para" id="p-274">The source code for this problem is shown here:</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. parsonsprob:: mt1dict1ex
:numbered: left

Complete the function greater_dictionary. Given a dictionary d and an integer cutoff, return a dictionary that contains only the key-value pairs where they key is greater than or equal to cutoff.
-----
def greater_dictionary(d, cutoff):
=====
def greater_dictionary(self, d, cutoff): #paired
=====
result = {}
=====
for key in d.keys():
=====
for key in range(d): #paired
=====
if key &gt;= cutoff:
=====
if key &gt; cutoff: #paired
=====
result[key] = d[key]
=====
d[key] = result[key] #paired
=====
return result
</code></pre></div></section><section class="subsection" id="writing-exercises_example-with-unpaired-distractor"><h4 class="heading hide-type">
<span class="type">Subsection</span> <span class="codenumber">8.3.4</span> <span class="title">Example with Unpaired Distractor</span>
</h4>
<div class="para" id="p-275">Here is an example with an unpaired distractor.</div>
<article class="exercise exercise-like" id="2_swapex"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.3.3</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone parsons_section" style="max-width: none;"><div data-component="parsons" class="parsons" id="2_swapex">
<div class="parsons_question parsons-text"><div class="para" id="p-276">The following has the correct code to ‘swap’ the values in x and y (so that x ends up with y’s initial value and y ends up with x’s initial value), but the code is mixed up and contains &lt;b&gt;one extra block&lt;/b&gt; which is not needed in a correct solution. Drag the needed blocks from the left into the correct order on the right. Check your solution by clicking on the &lt;i&gt;Check Me&lt;/i&gt; button. You will be told if any of the blocks are in the wrong order or if you need to remove one or more blocks.</div></div>
<pre class="parsonsblocks" data-question_label="" style="visibility: hidden;" data-language="python" data-noindent="true">int x = 3;
int y = 5;
int temp = 0;
---
temp = x;
---
x = y;
---
y = temp;
---
y = x;  #distractor</pre>
</div></div></div></article><div class="para" id="p-277">The source code for this problem is shown here.</div>
<div class="code-box" style="width: 100%; margin-left: 0%; margin-right: 0%;"><pre class="program"><code class="language-none">.. parsonsprob:: 2_swapex
:noindent:

The following has the correct code to 'swap' the values in x and y (so that x ends up with y's initial value and y ends up with x's initial value), but the code is mixed up and contains &lt;b&gt;one extra block&lt;/b&gt; which is not needed in a correct solution. Drag the needed blocks from the left into the correct order on the right. Check your solution by clicking on the &lt;i&gt;Check Me&lt;/i&gt; button. You will be told if any of the blocks are in the wrong order or if you need to remove one or more blocks.
-----
int x = 3;
int y = 5;
int temp = 0;
=====
temp = x;
=====
x = y;
=====
y = temp;
=====
y = x; #distractor
</code></pre></div></section></section><section class="section" id="writing-exercises_pro-tip-test-your-exercises-locally"><h3 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">8.4</span> <span class="title">Pro Tip – Test Your Exercises Locally</span>
</h3>
<div class="para" id="p-278">This is definitey an <dfn class="terminology">advanced topic</dfn>  If you are not comfortable with the command line you can definitely skip this section.  But if you are at least somewhat comfortable with the command line and an editor then this can save you a bunch of time!  The workflow is simple, edit your questions locally, then paste them into the interface once you have them working and debugged.</div>
<div class="para" id="p-279">Here is a short video that demonstrates how to install the runestone command in under two minutes.</div>
<blockquote class="blockquote" id="blockquote-2"><div class="video-box" style="width: 100%;padding-top: 56.25%; margin-left: 0%; margin-right: 0%;"><iframe id="QWX1eIpDBe0" class="video" allowfullscreen="" src="https://www.youtube-nocookie.com/embed/QWX1eIpDBe0?&amp;modestbranding=1&amp;rel=0"></iframe></div></blockquote>
<div class="para" id="p-280">Here are the steps I go through in the video.  This assumes you already have python installed on your system.</div>
<div class="para logical" id="p-281"><ol class="decimal">
<li id="li-128"><div class="para" id="p-282">
<code class="code-inline tex2jax_ignore">virtualenv runestone</code>  - Creates a python virtual environment so you don’t need administrator privileges to install the rest of what you need.  If virtualenv is not installed you should do <code class="code-inline tex2jax_ignore">pip install virtualenv</code> (Python 2.7.x) or <code class="code-inline tex2jax_ignore">python -m venv runestone</code> (Python &gt; 3.5).  If the pip command is not found then you should upgrade your version of Python to something newer.  All modern version of Python come with pip pre-installed.</div></li>
<li id="li-129"><div class="para" id="p-283">
<code class="code-inline tex2jax_ignore">. runestone/bin/activate</code> – this activates the virtual environment you created in the previous step.</div></li>
<li id="li-130"><div class="para" id="p-284">
<code class="code-inline tex2jax_ignore">mkdir assignments</code> – Create an assignment to do your runestone work in.</div></li>
<li id="li-131"><div class="para" id="p-285">
<code class="code-inline tex2jax_ignore">cd assignments</code>  – Make assignments your working directory.</div></li>
<li id="li-132"><div class="para" id="p-286">
<code class="code-inline tex2jax_ignore">runestone init</code> – Create a new project with some template directives</div></li>
<li id="li-133"><div class="para" id="p-287">
<code class="code-inline tex2jax_ignore">runestone build</code> – Convert the restructuredText to html</div></li>
<li id="li-134"><div class="para" id="p-288">
<code class="code-inline tex2jax_ignore">runestone serve</code> – start a simple web server so you can access your new project in your browser at <code class="code-inline tex2jax_ignore">http://localhost:8000/index.html</code>
</div></li>
</ol></div>
<div class="para" id="p-289">Once the works you can simply edit the index.rst file created by runestone init, and run <code class="code-inline tex2jax_ignore">runestone build</code> again to check your work.</div></section><section class="section" id="writing-exercises_add-your-own-question"><h3 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">8.5</span> <span class="title">Add your own Question</span>
</h3>
<div class="para" id="p-290">In the final exercise for the night, add your own question to the exercise bank using the web interface.</div>
<div class="para" id="p-291">You can make it any kind of question you would like, but it would be great if it was a real exercise that you like to use in one of your own courses.</div>
<div class="para" id="p-292">To keep the build system from bogging down, you can let me know once you have saved an exercise and I’ll pull it into the sigcse2019 book and build it for you.</div>
<article class="warning remark-like" id="warning-2"><h4 class="heading">
<span class="type">Warning</span><span class="space"> </span><span class="codenumber">8.5.1</span><span class="period">.</span>
</h4>
<div class="para" id="p-293">Make sure you assign at least 0 points to your question before you press done.  If everything works, you will get a popup telling you that the question has been added successfully.</div></article></section><section class="exercises" id="writing-exercises_exercises"><h3 class="heading hide-type">
<span class="type">Exercises</span> <span class="codenumber">8.6</span> <span class="title">Exercises</span>
</h3>
<div class="para" id="p-294">this page left intentionally blank</div></section></section><div class="hidden-content tex2jax_ignore" id="hk-fn-9"><div class="fn"><code class="code-inline tex2jax_ignore">https://docs.python.org/2/library/unittest.html#assert-methods</code></div></div>
</div>
<div class="ptx-content-footer">
<a class="previous-button button" href="peer_instruction.html" title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="top-button button" href="#" title="Top"><span class="icon">^</span><span class="name">Top</span></a><a class="next-button button" href="backmatter-1.html" title="Next"><span class="name">Next</span><span class="icon">&gt;</span></a>
</div></main>
</div>
<div id="ptx-page-footer" class="ptx-page-footer">
<a class="pretext-link" href="https://pretextbook.org" title="PreTeXt"><div class="logo"><svg xmlns="http://www.w3.org/2000/svg" width="100%" height="100%" viewBox="338 3000 8772 6866"><g style="stroke-width:.025in; stroke:black; fill:none"><polyline points="472,3590 472,9732 " style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round; "></polyline><path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;" d="M 4724,9448 A 4660 4660  0  0  1  8598  9259"></path><path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;" d="M 4488,9685 A 4228 4228  0  0  0   472  9732"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;" d="M 4724,3590 A 4241 4241  0  0  1  8598  3496"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:round;" d="M 850,3496  A 4241 4241  0  0  1  4724  3590"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:round;" d="M 850,9259  A 4507 4507  0  0  1  4724  9448"></path><polyline points="5385,4299 4062,8125" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8598,3496 8598,9259" style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="850,3496 850,9259" style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="4960,9685 4488,9685" style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="3070,4582 1889,6141 3070,7700" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="6418,4582 7600,6141 6418,7700" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8976,3590 8976,9732" style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;" d="M 4960,9685 A 4228 4228  0  0  1  8976  9732"></path></g></svg></div></a><a class="runestone-link" href="https://runestone.academy" title="Runestone Academy"><img class="logo" src="https://runestone.academy/runestone/static/images/RAIcon_cropped.png"></a><a class="mathjax-link" href="https://www.mathjax.org" title="MathJax"><img class="logo" src="https://www.mathjax.org/badge/badge-square-2.png"></a>
</div>
</body>
</html>
