<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<h5 class="heading"><span class="type">Paragraph</span></h5>
<div class="para">If you are not familiar with Python unittests they are pretty easy to write. You create your own class that is a subclass of TestCase, or in our work TestCaseGui so we get some graphical output. Your tests are all methods of the class and must start with the word “test”. There are a host of assertXXXX functions that you can use. Check out the <a class="external" href="https://docs.python.org/2/library/unittest.html#assert-methods" target="_blank">unittest documentation here</a><a href="" class="fn-knowl original" data-knowl="./knowl/fn-25-hidden.html" title="Footnote 8.1.1"><sup> 1 </sup></a>
</div>
<span class="incontext"><a href="chapter-8.html#p-338" class="internal">in-context</a></span>
</body>
</html>
