<!DOCTYPE html>
<html lang="en-US" dir="ltr">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h5 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.1.4</span><span class="period">.</span>
</h5>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="sigcse_ex4">
<div class="ac_question" id="sigcse_ex4_question"><div class="para" id="p-270">Write a Python function to sum the first N numbers starting with 0. So if N is 4 then your function should add 0 + 1 + 2 + 3</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="sigcse_ex4_editor" data-question_label="" data-codelens="true">def sum_first_n(N):
# your code here

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(sum_first_n(4),6,feedback="0 + 1 + 2 + 3 == 6")
self.assertEqual(sum_first_n(0),0,feedback="summing 0 numbers should be 0")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="chapter-8.html#sigcse_ex4" class="internal">in-context</a></span>
</body>
</html>
