<?xml version="1.0"?>
<section xml:id="writing-exercises_adding-an-exercise-to-the-exercise-bank">
  <title>Adding an Exercise to the Exercise Bank</title>
  <subsection xml:id="writing_exer_intro">
   <title>Create New Exercises</title>

    <video xml:id="create-new-exercises-video" youtube="IPvW9oKg1eA" width="auto" />
    
    <p>The new Assignment Builder makes it easy to write your own exercises! Although all of the Runestone books have a pretty good selection of exercises, you can never have too many exercises. We hope to crowd source an enormous collection of questions and problems in our Exercise database. You can help by adding custom exercises for your own course to the database! Once your exercise is in the database, it is available for others to find using the search feature on the create assignment page. Community contributed exercises are shown with a person icon instead of a book icon under the source column.</p>

    <p>To create new exercises, start or go to any assignment in the Instructor Dashboard "Assignment Builder". In the Exercises section of the assignment, click on "Add Exercise" and select "+ Create New Exercise". </p>
    <figure align="center" xml:id="create_new_exercise_fig">
      <caption>Create new exercise</caption>
      <image source="Figures/createNewExercise.png" width="100%" alt="Create new exercise"/>
    </figure>
    
<p>There are many different types of exercises that you can create as shown below. Some of these are discussed further in the sections below.</p>
<figure align="center" xml:id="exercise_types_fig">
  <caption>Exercise Types</caption>   
  <image source="Figures/exerciseTypes.png" width="100%" alt="Exercise Types"/>
</figure>

 <p>Here are some examples of these different types of exercises. In the next pages, we will go through the steps to create some of these exercises.</p>  

  <p>Here is an example Multiple Choice Exercise:</p>
  <exercise label="java_mcq_1" xml:id="java_mcq_1">
    <statement>
      <p>What does the following code print?</p>
      <program language="java">
        <code>
System.out.println(2 + 3 * 5 - 1);
        </code>
      </program>
    </statement>

    <choices>
      <choice>
        <statement>
          <p>24</p>
        </statement>

        <feedback>
          <p>
            This would be true if it was System.out.println(((2 + 3) * 5) - 1),
            but without the parentheses the multiplication is done first.
          </p>
        </feedback>
      </choice>

      <choice>
        <statement>
          <p>14</p>
        </statement>

        <feedback>
          <p>
            This would be true if it was System.out.println(2 + (3 * (5 - 1))),
            but without the parentheses the multiplication is done first and the
            addition and subtraction are handled from left to right.
          </p>
        </feedback>
      </choice>

      <choice>
        <statement>
          <p>This will give a compile time error.</p>
        </statement>

        <feedback>
          <p>
            This will compile and run. Try it in DrJava. Look up operator
            precedence in Java.
          </p>
        </feedback>
      </choice>

      <choice correct="yes">
        <statement>
          <p>16</p>
        </statement>

        <feedback>
          <p>
            Correct! Remember PEMDAS. The multiplication is done first (3 * 5 = 15) and then the addition (2 + 15 = 17) and finally the subtraction (17 - 1 = 16).
          </p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  
  <p>Parsons (Mixed up Code) Problem:</p>

   <exercise label="mixed-hi-paired" xml:id="mixed-hi-paired" numbered="yes" adaptive="yes" indentation="hide" language="java">
      <statement>
        <p>
          The following has all the correct Java code to print out “Hi there!” when
          the code is run, but the code is mixed up and contains some extra
          blocks with errors. Drag the needed blocks from left to right and put
          them in the correct order. Click on the “Check Me” button to check
          your solution.
        </p>
      </statement>

      <blocks>
        <block order="6">
          <choice correct="yes">
            <cline>public class Hello</cline>
            <cline>{</cline>
          </choice>

          <choice>
            <cline>public Class Hello</cline>
            <cline>{</cline>
          </choice>
        </block>

        <block order="5">
          <choice correct="yes">
            <cline>   public static void main(String[] args)</cline>
            <cline>   {</cline>
          </choice>

          <choice>
            <cline>   public static void main()</cline>
            <cline>   {</cline>
          </choice>
        </block>

        <block order="8">
          <choice correct="yes">
            <cline>      System.out.println("Hi there!");</cline>
          </choice>

          <choice>
            <cline>      System.out.println("Hi there!")</cline>
            <cline></cline>
          </choice>
        </block>

        <block order="2">
          <cline>   }</cline>
          <cline>}</cline>
        </block>
      </blocks>
    </exercise>

 
   <p>Here is an example of a simple Active Code exercise in Python:</p>
    <exercise label="active_code_example">
      <statement>
        <p>Write a Python function that takes two numbers as input and returns their sum.</p>
      </statement>
      <program interactive="activecode" language="python">
        <code>
def add(a, b):
  # complete this function to add a and b and return the sum
        </code>
        <tests>
from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
      self.assertEqual(add(2,2),4,"Call to add(2,2) should return 4")
      self.assertAlmostEqual(add(2.0,3.0), 5.0, 1, "Call to add(2.0,3.0) should return 5.0")

myTests().main()
        </tests>
      </program>
    </exercise>

  <p>Here is an example of a Drag and Drop matching questions. You can also have Matching questions where
  you draw lines to match many to many objects.  </p>
    <exercise label="java_matching">
      <statement>
        <p>
          Drag the definition from the left and drop it on the correct concept
          on the right. Click the "Check Me" button to see if you are correct.
        </p>
      </statement>

      <feedback>
        <p>Review the summaries above.</p>
      </feedback>

      <cardsort>
        <match>
          <premise order="1">a specific instance of a class with defined attributes</premise>
          <response>object</response>
        </match>

        <match>
          <premise order="2">defines a new data type with instance variables and methods</premise>
          <response>class</response>
        </match>

        <match>
          <premise order="3">define the behaviors or functions for objects</premise>
          <response>methods</response>
        </match>

        <match>
          <premise order="4">
            code that is used to create new objects and initialize the object's attributes
          </premise>

          <response>constructors</response>
        </match>

        <match>
          <premise order="5">the values or data passed to an object's method</premise>
          <response>parameters</response>
        </match>
      </cardsort>
    </exercise>

 <p>Clickable Area:</p>
  <exercise label="date_constructor">
      <statement>
        <p>Click on the constructor headers (signatures)</p>
      </statement>

      <feedback>
        <p>
          Constructors are public and have the same name as the class. Click on
          the constructor headers which are the first line of the constructors
          showing their name and parameters.
        </p>
      </feedback>

      <areas>
        <cline><area correct="no">public class Date {</area></cline>
        <cline></cline>
        <cline>    <area correct="no">private int year;</area></cline>
        <cline>    <area correct="no">private int month;</area></cline>
        <cline>    <area correct="no">private int day;</area></cline>
        <cline></cline>
        <cline>    <area correct="yes">public Date() </area></cline>
        <cline>        <area correct="no">{ /** Implementation not shown */ }</area></cline>
        <cline></cline>
        <cline>    <area correct="yes">public Date(int year, int month, int day) </area></cline>
        <cline>        <area correct="no">{ /** Implementation not shown */ }</area></cline>
        <cline></cline>
        <cline>     <area correct="no">public void print() </area></cline>
        <cline>        <area correct="no">{ /** Implementation not shown */ }</area></cline>
        <cline></cline>
        <cline><area correct="no">}</area></cline>
      </areas>
    </exercise>

  
 


  <p>Fill in the Blank:</p>
  <exercise label="fillDecVar2">
      <statement>
        <p>What Java data type should you use for a shoe size like 8.5?</p>
        <p>
          <fillin answer="double" mode="string" />
        </p>
      </statement>

      <evaluation>
        <evaluate>
          <test>
            <strcmp use-answer="yes" />
            <feedback>
              Correct.  Any variable that needs to values after the decimal point should be declared as a double.
            </feedback>
          </test>

          <test>
            <!-- else -->
            <feedback>What type allows for a decimal value?</feedback>
          </test>
        </evaluate>
      </evaluation>
    </exercise>
  <p>Short Answer</p>
   <exercise label="pbj">
      <statement>
        <p>
        Write an algorithm for someone (maybe a robot) to make a peanut butter and jelly sandwich.  
        Be sure to include at least 5 precise steps in order. 
        Have someone (maybe your instructor) act it out. Were your instructions precise enough?  
        </p>
      </statement>

      <response />
    </exercise>


  <p>Poll</p>

    <query label="qprogramming">
      <statement>
        <p>I am a</p>
      </statement>

      <choices>
        <choice>
          <p>beginner programmer</p>
        </choice>

        <choice>
          <p>intermediate programmer</p>
        </choice>

        <choice>
          <p>expert programmer</p>
        </choice>

        <choice>
          <p>prefer not to answer</p>
        </choice>
      </choices>
    </query>

  <p>Select Questions allow the students to choose between already existing questions:</p>
<exercise label="toggle1">
  <select grade="first" questions="mixed-hi-paired java_mcq_1"/>
</exercise>
  
  </subsection>
</section>